---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Ex 1.1

```yaml
type: NormalExercise 
lang: sql
xp: 100 
skills: 1
key: 4629420721   
```


Do some data science.


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}
SELECT * FROM
```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Case Study: Data Driven Decision Making with SQL

```yaml
type: VideoExercise 
xp: 50 
key: ba54eb3e52   
```

`@projector_key`
1c1e526409e0f865faeabdff8c9fba19
